


<!DOCTYPE html>
<html lang="en" class="">
  <head prefix="og: http://ogp.me/ns# fb: http://ogp.me/ns/fb# object: http://ogp.me/ns/object# article: http://ogp.me/ns/article# profile: http://ogp.me/ns/profile#">
    <meta charset='utf-8'>
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta http-equiv="Content-Language" content="en">
    
    
    <title>datacamp/README.md at master · marcionicolau/datacamp · GitHub</title>
    <link rel="search" type="application/opensearchdescription+xml" href="/opensearch.xml" title="GitHub">
    <link rel="fluid-icon" href="https://github.com/fluidicon.png" title="GitHub">
    <link rel="apple-touch-icon" sizes="57x57" href="/apple-touch-icon-114.png">
    <link rel="apple-touch-icon" sizes="114x114" href="/apple-touch-icon-114.png">
    <link rel="apple-touch-icon" sizes="72x72" href="/apple-touch-icon-144.png">
    <link rel="apple-touch-icon" sizes="144x144" href="/apple-touch-icon-144.png">
    <meta property="fb:app_id" content="1401488693436528">

      <meta content="@github" name="twitter:site" /><meta content="summary" name="twitter:card" /><meta content="marcionicolau/datacamp" name="twitter:title" /><meta content="R interface to create interactive courses for www.datacamp.com" name="twitter:description" /><meta content="https://avatars1.githubusercontent.com/u/1499733?v=3&amp;s=400" name="twitter:image:src" />
      <meta content="GitHub" property="og:site_name" /><meta content="object" property="og:type" /><meta content="https://avatars1.githubusercontent.com/u/1499733?v=3&amp;s=400" property="og:image" /><meta content="marcionicolau/datacamp" property="og:title" /><meta content="https://github.com/marcionicolau/datacamp" property="og:url" /><meta content="R interface to create interactive courses for www.datacamp.com" property="og:description" />
      <meta name="browser-stats-url" content="/_stats">
    <link rel="assets" href="https://assets-cdn.github.com/">
    
    <meta name="pjax-timeout" content="1000">
    

    <meta name="msapplication-TileImage" content="/windows-tile.png">
    <meta name="msapplication-TileColor" content="#ffffff">
    <meta name="selected-link" value="repo_source" data-pjax-transient>
      <meta name="google-analytics" content="UA-3769691-2">

    <meta content="collector.githubapp.com" name="octolytics-host" /><meta content="collector-cdn.github.com" name="octolytics-script-host" /><meta content="github" name="octolytics-app-id" /><meta content="C1BEC102:4D96:73B0DE7:55104B18" name="octolytics-dimension-request_id" />
    
    <meta content="Rails, view, blob#show" name="analytics-event" />

    
    <link rel="icon" type="image/x-icon" href="https://assets-cdn.github.com/favicon.ico">


    <meta content="authenticity_token" name="csrf-param" />
<meta content="13YL9vPJXby//Fj3CX7ujHXj8wqHPHBN7inSWfke9oBf3pC+QdigEZPUtd1vYDmxH8MX8Xr+m3pEZot+uTqJoA==" name="csrf-token" />

    <link href="https://assets-cdn.github.com/assets/github-099e0ecc2851c8aca89ef6dafa191df3b0f2a2c8ad34e134c5473ca1ba0a59b2.css" media="all" rel="stylesheet" />
    <link href="https://assets-cdn.github.com/assets/github2-1171344316fc088255ee2a06c271d14240f1a4e06985fe9e897762947872e858.css" media="all" rel="stylesheet" />
    
    


    <meta http-equiv="x-pjax-version" content="c0f32272c66bfb10ed7d46b7c88c6299">

      
  <meta name="description" content="R interface to create interactive courses for www.datacamp.com">
  <meta name="go-import" content="github.com/marcionicolau/datacamp git https://github.com/marcionicolau/datacamp.git">

  <meta content="1499733" name="octolytics-dimension-user_id" /><meta content="marcionicolau" name="octolytics-dimension-user_login" /><meta content="19275541" name="octolytics-dimension-repository_id" /><meta content="marcionicolau/datacamp" name="octolytics-dimension-repository_nwo" /><meta content="true" name="octolytics-dimension-repository_public" /><meta content="true" name="octolytics-dimension-repository_is_fork" /><meta content="15495987" name="octolytics-dimension-repository_parent_id" /><meta content="Data-Camp/datacamp" name="octolytics-dimension-repository_parent_nwo" /><meta content="15495987" name="octolytics-dimension-repository_network_root_id" /><meta content="Data-Camp/datacamp" name="octolytics-dimension-repository_network_root_nwo" />
  <link href="https://github.com/marcionicolau/datacamp/commits/master.atom" rel="alternate" title="Recent Commits to datacamp:master" type="application/atom+xml">

  </head>


  <body class="logged_out  env-production  vis-public fork page-blob">
    <a href="#start-of-content" tabindex="1" class="accessibility-aid js-skip-to-content">Skip to content</a>
    <div class="wrapper">
      
      
      


        
        <div class="header header-logged-out" role="banner">
  <div class="container clearfix">

    <a class="header-logo-wordmark" href="https://github.com/" data-ga-click="(Logged out) Header, go to homepage, icon:logo-wordmark">
      <span class="mega-octicon octicon-logo-github"></span>
    </a>

    <div class="header-actions" role="navigation">
        <a class="btn btn-primary" href="/join" data-ga-click="(Logged out) Header, clicked Sign up, text:sign-up">Sign up</a>
      <a class="btn" href="/login?return_to=%2Fmarcionicolau%2Fdatacamp%2Fblob%2Fmaster%2FREADME.md" data-ga-click="(Logged out) Header, clicked Sign in, text:sign-in">Sign in</a>
    </div>

    <div class="site-search repo-scope js-site-search" role="search">
      <form accept-charset="UTF-8" action="/marcionicolau/datacamp/search" class="js-site-search-form" data-global-search-url="/search" data-repo-search-url="/marcionicolau/datacamp/search" method="get"><div style="margin:0;padding:0;display:inline"><input name="utf8" type="hidden" value="&#x2713;" /></div>
  <input type="text"
    class="js-site-search-field is-clearable"
    data-hotkey="s"
    name="q"
    placeholder="Search"
    data-global-scope-placeholder="Search GitHub"
    data-repo-scope-placeholder="Search"
    tabindex="1"
    autocapitalize="off">
  <div class="scope-badge">This repository</div>
</form>
    </div>

      <ul class="header-nav left" role="navigation">
          <li class="header-nav-item">
            <a class="header-nav-link" href="/explore" data-ga-click="(Logged out) Header, go to explore, text:explore">Explore</a>
          </li>
          <li class="header-nav-item">
            <a class="header-nav-link" href="/features" data-ga-click="(Logged out) Header, go to features, text:features">Features</a>
          </li>
          <li class="header-nav-item">
            <a class="header-nav-link" href="https://enterprise.github.com/" data-ga-click="(Logged out) Header, go to enterprise, text:enterprise">Enterprise</a>
          </li>
          <li class="header-nav-item">
            <a class="header-nav-link" href="/blog" data-ga-click="(Logged out) Header, go to blog, text:blog">Blog</a>
          </li>
      </ul>

  </div>
</div>



      <div id="start-of-content" class="accessibility-aid"></div>
          <div class="site" itemscope itemtype="http://schema.org/WebPage">
    <div id="js-flash-container">
      
    </div>
    <div class="pagehead repohead instapaper_ignore readability-menu">
      <div class="container">
        
<ul class="pagehead-actions">

  <li>
      <a href="/login?return_to=%2Fmarcionicolau%2Fdatacamp"
    class="btn btn-sm btn-with-count tooltipped tooltipped-n"
    aria-label="You must be signed in to watch a repository" rel="nofollow">
    <span class="octicon octicon-eye"></span>
    Watch
  </a>
  <a class="social-count" href="/marcionicolau/datacamp/watchers">
    1
  </a>

  </li>

  <li>
      <a href="/login?return_to=%2Fmarcionicolau%2Fdatacamp"
    class="btn btn-sm btn-with-count tooltipped tooltipped-n"
    aria-label="You must be signed in to star a repository" rel="nofollow">
    <span class="octicon octicon-star"></span>
    Star
  </a>

    <a class="social-count js-social-count" href="/marcionicolau/datacamp/stargazers">
      0
    </a>

  </li>

    <li>
      <a href="/login?return_to=%2Fmarcionicolau%2Fdatacamp"
        class="btn btn-sm btn-with-count tooltipped tooltipped-n"
        aria-label="You must be signed in to fork a repository" rel="nofollow">
        <span class="octicon octicon-repo-forked"></span>
        Fork
      </a>
      <a href="/marcionicolau/datacamp/network" class="social-count">
        12
      </a>
    </li>
</ul>

        <h1 itemscope itemtype="http://data-vocabulary.org/Breadcrumb" class="entry-title public">
          <span class="mega-octicon octicon-repo-forked"></span>
          <span class="author"><a href="/marcionicolau" class="url fn" itemprop="url" rel="author"><span itemprop="title">marcionicolau</span></a></span><!--
       --><span class="path-divider">/</span><!--
       --><strong><a href="/marcionicolau/datacamp" class="js-current-repository" data-pjax="#js-repo-pjax-container">datacamp</a></strong>

          <span class="page-context-loader">
            <img alt="" height="16" src="https://assets-cdn.github.com/assets/spinners/octocat-spinner-32-e513294efa576953719e4e2de888dd9cf929b7d62ed8d05f25e731d02452ab6c.gif" width="16" />
          </span>

            <span class="fork-flag">
              <span class="text">forked from <a href="/Data-Camp/datacamp">Data-Camp/datacamp</a></span>
            </span>
        </h1>
      </div><!-- /.container -->
    </div><!-- /.repohead -->

    <div class="container">
      <div class="repository-with-sidebar repo-container new-discussion-timeline  ">
        <div class="repository-sidebar clearfix">
            
<nav class="sunken-menu repo-nav js-repo-nav js-sidenav-container-pjax js-octicon-loaders"
     role="navigation"
     data-pjax="#js-repo-pjax-container"
     data-issue-count-url="/marcionicolau/datacamp/issues/counts">
  <ul class="sunken-menu-group">
    <li class="tooltipped tooltipped-w" aria-label="Code">
      <a href="/marcionicolau/datacamp" aria-label="Code" class="selected js-selected-navigation-item sunken-menu-item" data-hotkey="g c" data-selected-links="repo_source repo_downloads repo_commits repo_releases repo_tags repo_branches /marcionicolau/datacamp">
        <span class="octicon octicon-code"></span> <span class="full-word">Code</span>
        <img alt="" class="mini-loader" height="16" src="https://assets-cdn.github.com/assets/spinners/octocat-spinner-32-e513294efa576953719e4e2de888dd9cf929b7d62ed8d05f25e731d02452ab6c.gif" width="16" />
</a>    </li>


    <li class="tooltipped tooltipped-w" aria-label="Pull requests">
      <a href="/marcionicolau/datacamp/pulls" aria-label="Pull requests" class="js-selected-navigation-item sunken-menu-item" data-hotkey="g p" data-selected-links="repo_pulls /marcionicolau/datacamp/pulls">
          <span class="octicon octicon-git-pull-request"></span> <span class="full-word">Pull requests</span>
          <span class="js-pull-replace-counter"></span>
          <img alt="" class="mini-loader" height="16" src="https://assets-cdn.github.com/assets/spinners/octocat-spinner-32-e513294efa576953719e4e2de888dd9cf929b7d62ed8d05f25e731d02452ab6c.gif" width="16" />
</a>    </li>


  </ul>
  <div class="sunken-menu-separator"></div>
  <ul class="sunken-menu-group">

    <li class="tooltipped tooltipped-w" aria-label="Pulse">
      <a href="/marcionicolau/datacamp/pulse" aria-label="Pulse" class="js-selected-navigation-item sunken-menu-item" data-selected-links="pulse /marcionicolau/datacamp/pulse">
        <span class="octicon octicon-pulse"></span> <span class="full-word">Pulse</span>
        <img alt="" class="mini-loader" height="16" src="https://assets-cdn.github.com/assets/spinners/octocat-spinner-32-e513294efa576953719e4e2de888dd9cf929b7d62ed8d05f25e731d02452ab6c.gif" width="16" />
</a>    </li>

    <li class="tooltipped tooltipped-w" aria-label="Graphs">
      <a href="/marcionicolau/datacamp/graphs" aria-label="Graphs" class="js-selected-navigation-item sunken-menu-item" data-selected-links="repo_graphs repo_contributors /marcionicolau/datacamp/graphs">
        <span class="octicon octicon-graph"></span> <span class="full-word">Graphs</span>
        <img alt="" class="mini-loader" height="16" src="https://assets-cdn.github.com/assets/spinners/octocat-spinner-32-e513294efa576953719e4e2de888dd9cf929b7d62ed8d05f25e731d02452ab6c.gif" width="16" />
</a>    </li>
  </ul>


</nav>

              <div class="only-with-full-nav">
                  
<div class="clone-url open"
  data-protocol-type="http"
  data-url="/users/set_protocol?protocol_selector=http&amp;protocol_type=clone">
  <h3><span class="text-emphasized">HTTPS</span> clone URL</h3>
  <div class="input-group js-zeroclipboard-container">
    <input type="text" class="input-mini input-monospace js-url-field js-zeroclipboard-target"
           value="https://github.com/marcionicolau/datacamp.git" readonly="readonly">
    <span class="input-group-button">
      <button aria-label="Copy to clipboard" class="js-zeroclipboard btn btn-sm zeroclipboard-button" data-copied-hint="Copied!" type="button"><span class="octicon octicon-clippy"></span></button>
    </span>
  </div>
</div>

  
<div class="clone-url "
  data-protocol-type="subversion"
  data-url="/users/set_protocol?protocol_selector=subversion&amp;protocol_type=clone">
  <h3><span class="text-emphasized">Subversion</span> checkout URL</h3>
  <div class="input-group js-zeroclipboard-container">
    <input type="text" class="input-mini input-monospace js-url-field js-zeroclipboard-target"
           value="https://github.com/marcionicolau/datacamp" readonly="readonly">
    <span class="input-group-button">
      <button aria-label="Copy to clipboard" class="js-zeroclipboard btn btn-sm zeroclipboard-button" data-copied-hint="Copied!" type="button"><span class="octicon octicon-clippy"></span></button>
    </span>
  </div>
</div>



<p class="clone-options">You can clone with
  <a href="#" class="js-clone-selector" data-protocol="http">HTTPS</a> or <a href="#" class="js-clone-selector" data-protocol="subversion">Subversion</a>.
  <a href="https://help.github.com/articles/which-remote-url-should-i-use" class="help tooltipped tooltipped-n" aria-label="Get help on which URL is right for you.">
    <span class="octicon octicon-question"></span>
  </a>
</p>



                <a href="/marcionicolau/datacamp/archive/master.zip"
                   class="btn btn-sm sidebar-button"
                   aria-label="Download the contents of marcionicolau/datacamp as a zip file"
                   title="Download the contents of marcionicolau/datacamp as a zip file"
                   rel="nofollow">
                  <span class="octicon octicon-cloud-download"></span>
                  Download ZIP
                </a>
              </div>
        </div><!-- /.repository-sidebar -->

        <div id="js-repo-pjax-container" class="repository-content context-loader-container" data-pjax-container>
          

<a href="/marcionicolau/datacamp/blob/0874ed71b6315ec4a86e97579180d4df8b74460a/README.md" class="hidden js-permalink-shortcut" data-hotkey="y">Permalink</a>

<!-- blob contrib key: blob_contributors:v21:dc4ed2b774b65becd1d9481d5761d62b -->

<div class="file-navigation js-zeroclipboard-container">
  
<div class="select-menu js-menu-container js-select-menu left">
  <span class="btn btn-sm select-menu-button js-menu-target css-truncate" data-hotkey="w"
    data-master-branch="master"
    data-ref="master"
    title="master"
    role="button" aria-label="Switch branches or tags" tabindex="0" aria-haspopup="true">
    <span class="octicon octicon-git-branch"></span>
    <i>branch:</i>
    <span class="js-select-button css-truncate-target">master</span>
  </span>

  <div class="select-menu-modal-holder js-menu-content js-navigation-container" data-pjax aria-hidden="true">

    <div class="select-menu-modal">
      <div class="select-menu-header">
        <span class="select-menu-title">Switch branches/tags</span>
        <span class="octicon octicon-x js-menu-close" role="button" aria-label="Close"></span>
      </div>

      <div class="select-menu-filters">
        <div class="select-menu-text-filter">
          <input type="text" aria-label="Filter branches/tags" id="context-commitish-filter-field" class="js-filterable-field js-navigation-enable" placeholder="Filter branches/tags">
        </div>
        <div class="select-menu-tabs">
          <ul>
            <li class="select-menu-tab">
              <a href="#" data-tab-filter="branches" data-filter-placeholder="Filter branches/tags" class="js-select-menu-tab">Branches</a>
            </li>
            <li class="select-menu-tab">
              <a href="#" data-tab-filter="tags" data-filter-placeholder="Find a tag…" class="js-select-menu-tab">Tags</a>
            </li>
          </ul>
        </div>
      </div>

      <div class="select-menu-list select-menu-tab-bucket js-select-menu-tab-bucket" data-tab-filter="branches">

        <div data-filterable-for="context-commitish-filter-field" data-filterable-type="substring">


            <a class="select-menu-item js-navigation-item js-navigation-open "
               href="/marcionicolau/datacamp/blob/develop/README.md"
               data-name="develop"
               data-skip-pjax="true"
               rel="nofollow">
              <span class="select-menu-item-icon octicon octicon-check"></span>
              <span class="select-menu-item-text css-truncate-target" title="develop">
                develop
              </span>
            </a>
            <a class="select-menu-item js-navigation-item js-navigation-open selected"
               href="/marcionicolau/datacamp/blob/master/README.md"
               data-name="master"
               data-skip-pjax="true"
               rel="nofollow">
              <span class="select-menu-item-icon octicon octicon-check"></span>
              <span class="select-menu-item-text css-truncate-target" title="master">
                master
              </span>
            </a>
        </div>

          <div class="select-menu-no-results">Nothing to show</div>
      </div>

      <div class="select-menu-list select-menu-tab-bucket js-select-menu-tab-bucket" data-tab-filter="tags">
        <div data-filterable-for="context-commitish-filter-field" data-filterable-type="substring">


            <div class="select-menu-item js-navigation-item ">
              <span class="select-menu-item-icon octicon octicon-check"></span>
              <a href="/marcionicolau/datacamp/tree/v0.2/README.md"
                 data-name="v0.2"
                 data-skip-pjax="true"
                 rel="nofollow"
                 class="js-navigation-open select-menu-item-text css-truncate-target"
                 title="v0.2">v0.2</a>
            </div>
            <div class="select-menu-item js-navigation-item ">
              <span class="select-menu-item-icon octicon octicon-check"></span>
              <a href="/marcionicolau/datacamp/tree/v0.1/README.md"
                 data-name="v0.1"
                 data-skip-pjax="true"
                 rel="nofollow"
                 class="js-navigation-open select-menu-item-text css-truncate-target"
                 title="v0.1">v0.1</a>
            </div>
        </div>

        <div class="select-menu-no-results">Nothing to show</div>
      </div>

    </div>
  </div>
</div>

  <div class="btn-group right">
    <a href="/marcionicolau/datacamp/find/master"
          class="js-show-file-finder btn btn-sm empty-icon tooltipped tooltipped-s"
          data-pjax
          data-hotkey="t"
          aria-label="Quickly jump between files">
      <span class="octicon octicon-list-unordered"></span>
    </a>
    <button aria-label="Copy file path to clipboard" class="js-zeroclipboard btn btn-sm zeroclipboard-button" data-copied-hint="Copied!" type="button"><span class="octicon octicon-clippy"></span></button>
  </div>

  <div class="breadcrumb js-zeroclipboard-target">
    <span class='repo-root js-repo-root'><span itemscope="" itemtype="http://data-vocabulary.org/Breadcrumb"><a href="/marcionicolau/datacamp" class="" data-branch="master" data-direction="back" data-pjax="true" itemscope="url"><span itemprop="title">datacamp</span></a></span></span><span class="separator">/</span><strong class="final-path">README.md</strong>
  </div>
</div>

<include-fragment class="commit commit-loader file-history-tease" src="/marcionicolau/datacamp/contributors/master/README.md">
  <div class="file-history-tease-header">
    Fetching contributors&hellip;
  </div>

  <div class="participation">
    <p class="loader-loading"><img alt="" height="16" src="https://assets-cdn.github.com/assets/spinners/octocat-spinner-32-EAF2F5-0bdc57d34b85c4a4de9d0d1db10cd70e8a95f33ff4f46c5a8c48b4bf4e5a9abe.gif" width="16" /></p>
    <p class="loader-error">Cannot retrieve contributors at this time</p>
  </div>
</include-fragment>
<div class="file">
  <div class="file-header">
    <div class="file-actions">

      <div class="btn-group">
        <a href="/marcionicolau/datacamp/raw/master/README.md" class="btn btn-sm " id="raw-url">Raw</a>
          <a href="/marcionicolau/datacamp/blame/master/README.md" class="btn btn-sm js-update-url-with-hash">Blame</a>
        <a href="/marcionicolau/datacamp/commits/master/README.md" class="btn btn-sm " rel="nofollow">History</a>
      </div>


          <button type="button" class="octicon-btn disabled tooltipped tooltipped-n" aria-label="You must be signed in to make or propose changes">
            <span class="octicon octicon-pencil"></span>
          </button>

        <button type="button" class="octicon-btn octicon-btn-danger disabled tooltipped tooltipped-n" aria-label="You must be signed in to make or propose changes">
          <span class="octicon octicon-trashcan"></span>
        </button>
    </div>

    <div class="file-info">
        219 lines (196 sloc)
        <span class="file-info-divider"></span>
      9.985 kb
    </div>
  </div>
    <div id="readme" class="blob instapaper_body">
    <article class="markdown-body entry-content" itemprop="mainContentOfPage"><p><a href="https://travis-ci.org/Data-Camp/datacamp"><img src="https://camo.githubusercontent.com/876609fd503907c4415796228cf0f7c72e23b97a/68747470733a2f2f6170692e7472617669732d63692e6f72672f446174612d43616d702f6461746163616d702e7376673f6272616e63683d6d6173746572" alt="Build Status" data-canonical-src="https://api.travis-ci.org/Data-Camp/datacamp.svg?branch=master" style="max-width:100%;"></a></p>

<p align="center">
<a href="https://camo.githubusercontent.com/47bf7627e5f381d09801b0cf98d4d7cc4d3b44cd/68747470733a2f2f73332e616d617a6f6e6177732e636f6d2f6173736574732e6461746163616d702e636f6d2f696d672f6c6f676f2f6c6f676f5f626c75655f66756c6c2e737667" target="_blank"><img src="https://camo.githubusercontent.com/47bf7627e5f381d09801b0cf98d4d7cc4d3b44cd/68747470733a2f2f73332e616d617a6f6e6177732e636f6d2f6173736574732e6461746163616d702e636f6d2f696d672f6c6f676f2f6c6f676f5f626c75655f66756c6c2e737667" width="250" data-canonical-src="https://s3.amazonaws.com/assets.datacamp.com/img/logo/logo_blue_full.svg" style="max-width:100%;"></a>
</p>

<h1>
<a id="user-content-build-your-own-interactive-courses" class="anchor" href="#build-your-own-interactive-courses" aria-hidden="true"><span class="octicon octicon-link"></span></a>Build your own interactive course(s)!</h1>

<h2>
<a id="user-content-introduction" class="anchor" href="#introduction" aria-hidden="true"><span class="octicon octicon-link"></span></a>Introduction</h2>

<p>Creating interactive courses for data analysis with R can and should be easy. This tutorial explains <strong>how you can create interactive courses</strong> on the DataCamp platform. Furthermore, the tutorial explains the <strong>Submission Correctness Test</strong> - the key concept to add interactivity to your course or tutorial.</p>

<h2>
<a id="user-content-creating-and-managing-interactive-courses" class="anchor" href="#creating-and-managing-interactive-courses" aria-hidden="true"><span class="octicon octicon-link"></span></a>Creating and managing interactive courses</h2>

<p>Courses are created locally using R Markdown files and managed with a complementary web interface on datacamp.com.</p>

<ul class="task-list">
<li>
<strong>The R interface - Use R Markdown and the datacamp R package to create interactive courses:</strong><br>
The datacamp R package enables you to write interactive R courses in R Markdown (e.g. using RStudio locally) and to then upload them to the datacamp.com platform. The DataCamp package is essentially a wrapper around the <a href="http://slidify.org/">slidify</a> package. This way, you can also use existing R Markdown files as the basis of your interactive courses and collaborate easily with others on courses (e.g. by sharing the Markdown files on GitHub). Have a look at the <a href="https://github.com/data-camp/introduction_to_R">R Markdown files</a> of the "Introduction to R" course.</li>
<li>
<strong>The web interface - Course management:</strong><br>
After uploading chapters to datacamp.com, you can use the <a href="https://teach.datacamp.com">webinterface</a> to view and to organize your courses with their chapters. You can use this interface to make sure that your course was uploaded correctly, to make various changes to the course's state and to get a preview on how the user will see the course.</li>
</ul>

<h2>
<a id="user-content-getting-started" class="anchor" href="#getting-started" aria-hidden="true"><span class="octicon octicon-link"></span></a>Getting started</h2>

<h3>
<a id="user-content-install-the-datacamp-r-package" class="anchor" href="#install-the-datacamp-r-package" aria-hidden="true"><span class="octicon octicon-link"></span></a>Install the datacamp R package</h3>

<div class="highlight highlight-ruby"><pre>install.packages(<span class="pl-s"><span class="pl-pds">"</span>devtools<span class="pl-pds">"</span></span>)
library(<span class="pl-s"><span class="pl-pds">"</span>devtools<span class="pl-pds">"</span></span>)
install_github(<span class="pl-s"><span class="pl-pds">"</span>datacampSCT<span class="pl-pds">"</span></span>, <span class="pl-s"><span class="pl-pds">"</span>data-camp<span class="pl-pds">"</span></span>)
install_github(<span class="pl-s"><span class="pl-pds">"</span>datacamp<span class="pl-pds">"</span></span>, <span class="pl-s"><span class="pl-pds">"</span>data-camp<span class="pl-pds">"</span></span>)
install_github(<span class="pl-s"><span class="pl-pds">"</span>slidify<span class="pl-pds">"</span></span>, <span class="pl-s"><span class="pl-pds">"</span>ramnathv<span class="pl-pds">"</span></span>, ref <span class="pl-k">=</span> <span class="pl-s"><span class="pl-pds">"</span>dev<span class="pl-pds">"</span></span>)
install_github(<span class="pl-s"><span class="pl-pds">"</span>slidifyLibraries<span class="pl-pds">"</span></span>, <span class="pl-s"><span class="pl-pds">"</span>ramnathv<span class="pl-pds">"</span></span>)
library(<span class="pl-s"><span class="pl-pds">"</span>datacamp<span class="pl-pds">"</span></span>)</pre></div>

<p>The code above will download and install the datacamp R package. Alternatively, you can also download the source code above.</p>

<h3>
<a id="user-content-scaffold-a-course-with-chapter" class="anchor" href="#scaffold-a-course-with-chapter" aria-hidden="true"><span class="octicon octicon-link"></span></a>Scaffold a course with chapter</h3>

<div class="highlight highlight-ruby"><pre>author_course(<span class="pl-s"><span class="pl-pds">"</span>course_name<span class="pl-pds">"</span></span>)</pre></div>

<p>The <code>author_course</code> function will:</p>

<ol class="task-list">
<li>create a folder in your current working directory with the name "course_name"</li>
<li>create a <code>course.yml</code> file, a scaffold to create your course</li>
<li>create a <code>chapter1.Rmd</code> file, a scaffold for creating your first chapter
The generated template contains extra instructions on the building blocks of chapters. We strongly advise you to put your folder under git source control at this point.</li>
</ol>

<h3>
<a id="user-content-directory-structure" class="anchor" href="#directory-structure" aria-hidden="true"><span class="octicon octicon-link"></span></a>Directory structure</h3>

<p>Every course consists of a <code>course.yml</code> file along with various <code>.Rmd</code> chapter files. An example course is structured as follows:</p>

<pre><code>.
|__ course.yml
|__ chapter1.Rmd
|__ chapter2.Rmd
|__ chapter3.Rmd
|__ ...
</code></pre>

<p>Note that the chapter <code>.Rmd</code> files can have any chosen name.</p>

<h3>
<a id="user-content-log-in-to-datacampcom" class="anchor" href="#log-in-to-datacampcom" aria-hidden="true"><span class="octicon octicon-link"></span></a>Log in to datacamp.com</h3>

<div class="highlight highlight-ruby"><pre>datacamp_login()</pre></div>

<p>This will prompt you for your username and password to log into the datacamp server. Note that you must also log into datacamp.com with your browser.</p>

<h3>
<a id="user-content-the-courseyml-file" class="anchor" href="#the-courseyml-file" aria-hidden="true"><span class="octicon octicon-link"></span></a>The <code>course.yml</code> file</h3>

<p>This file contains information about your course and the structure of its chapters. After running <code>author_course()</code>, a <code>course.yml</code> file is created in your course directory. This <code>course.yml</code> file has the following scaffold:</p>

<div class="highlight highlight-yml"><pre><span class="pl-c"># course.yml scaffold</span>
<span class="pl-s"><span class="pl-ent">title:</span> <span class="pl-s">Insert the course title here</span></span>
<span class="pl-s"><span class="pl-ent">author_field:</span> <span class="pl-s">The author(s) name(s)</span></span>
<span class="pl-s"><span class="pl-ent">description:</span> <span class="pl-s">Insert course description here</span></span></pre></div>

<p>Change the fields and call the function <code>upload_course()</code>. This will create your course and add its ID to the <code>course.yml</code> file. This way, your local course is linked to the one on datacamp.com.</p>

<div class="highlight highlight-yml"><pre><span class="pl-c"># course.yml after creation on server</span>
<span class="pl-c1"><span class="pl-ent">id:</span> 314</span>
<span class="pl-s"><span class="pl-ent">title:</span> <span class="pl-s">Hello world</span></span>
<span class="pl-s"><span class="pl-ent">author_field:</span> <span class="pl-s">Your name</span></span>
<span class="pl-s"><span class="pl-ent">description:</span> <span class="pl-s">A first course</span></span></pre></div>

<h3>
<a id="user-content-uploading-chapters-to-datacampcom" class="anchor" href="#uploading-chapters-to-datacampcom" aria-hidden="true"><span class="octicon octicon-link"></span></a>Uploading chapters to datacamp.com</h3>

<p>Every chapter is described in its corresponding <code>chapterX.Rmd</code> file. The first time that you create a chapter, it is a good idea to start from the <code>chapter1.Rmd</code> template that was generated using <code>author_course("course_name")</code>. This template contains information about the markup language, the different types of exercises, etc.</p>

<div class="highlight highlight-yml"><pre><span class="pl-c"># chapter1.Rmd scaffold</span>
<span class="pl-s">-<span class="pl-s">-- </span></span>
<span class="pl-s"><span class="pl-ent">title_meta  :</span> <span class="pl-s">Insert chapter preposition here</span></span>
<span class="pl-s"><span class="pl-ent">title       :</span> <span class="pl-s">Insert the chapter title here</span></span>
<span class="pl-s"><span class="pl-ent">description :</span> <span class="pl-s">What is this chapter about? Add here your description</span></span>
<span class="pl-s"><span class="pl-ent">framework   :</span> <span class="pl-s">datacamp</span></span>
<span class="pl-s"><span class="pl-ent">mode        :</span> <span class="pl-s">selfcontained</span></span>

<span class="pl-s">-<span class="pl-s">--</span></span>
<span class="pl-c">## The title of the first exercise</span>
...</pre></div>

<p>Adapt the scaffold and use the function <code>upload_chapter("chapter1.Rmd")</code> to add the chapter to your course on datacamp.com. This will create the chapter and add it to the <code>course.yml</code> file. Every time you add a new chapter to a course, it is added to the <code>course.yml</code> file. An example:</p>

<div class="highlight highlight-yml"><pre><span class="pl-c"># e.g. course.yml</span>
<span class="pl-c1"><span class="pl-ent">id:</span> 314</span>
<span class="pl-s"><span class="pl-ent">title:</span> <span class="pl-s">Hello world</span></span>
<span class="pl-s"><span class="pl-ent">author_field:</span> <span class="pl-s">Your name</span></span>
<span class="pl-s"><span class="pl-ent">description:</span> <span class="pl-s">A first course</span></span>
<span class="pl-s"><span class="pl-ent">chapters:</span></span>
  <span class="pl-c1">- <span class="pl-ent">chapter1.Rmd:</span> 753</span>
  <span class="pl-c1">- <span class="pl-ent">chapter2.Rmd:</span> 760</span></pre></div>

<p>As you can see, the file contains a list of chapters with a mapping of their file names on their ID. The order of the list dictates the order of the chapters within the course. This helps you to organize your course and to avoid overwriting chapters. For example, suppose we want chapter 2 to appear first on datacamp.com:</p>

<div class="highlight highlight-yml"><pre><span class="pl-c"># course.yml</span>
...
<span class="pl-s"><span class="pl-ent">chapters:</span></span>
  <span class="pl-c1">- <span class="pl-ent">chapter2.Rmd:</span> 760</span>
  <span class="pl-c1">- <span class="pl-ent">chapter1.Rmd:</span> 753</span></pre></div>

<p>Any changes to this list require the course to be uploaded with <code>upload_course()</code> for the changes to take effect.</p>

<h3>
<a id="user-content-share-the-love" class="anchor" href="#share-the-love" aria-hidden="true"><span class="octicon octicon-link"></span></a>Share the love</h3>

<p>A course is just a collection of chapters, and a chapter is essentially just a simple R Markdown document. Therefore, you can easily collaborate with others on a course. For example, by sharing the <a href="https://github.com/data-camp/introduction_to_R">R Markdown documents</a> on GitHub. We believe this "open-course" type of collaboration for interactive R and statistics courses will benefit everyone. Check out the open-source repo of our "Introduction to R" course and help us improve it.</p>

<h2>
<a id="user-content-submission-correctness-tests" class="anchor" href="#submission-correctness-tests" aria-hidden="true"><span class="octicon octicon-link"></span></a>Submission Correctness Tests</h2>

<p>The key ingredient to an interactive course is the Submission Correctness Test (SCT). For all information regarding SCTs, please refer to the <a href="https://github.com/Data-Camp/datacampSCT">datacampSCT package</a>.</p>

<h2>
<a id="user-content-most-frequently-asked-questions" class="anchor" href="#most-frequently-asked-questions" aria-hidden="true"><span class="octicon octicon-link"></span></a>Most frequently asked questions</h2>

<h5>
<a id="user-content-how-can-i-create-a-new-chapter" class="anchor" href="#how-can-i-create-a-new-chapter" aria-hidden="true"><span class="octicon octicon-link"></span></a>How can I create a new chapter?</h5>

<p>First, make sure that the current working directory matches with the course directory for which you want to create a chapter. Next, add a new chapter to the course with the help of the <code>author_chapter("chapter_name")</code> function. This will create and open an R Markdown file named <code>chapter_name.Rmd</code>.  </p>

<pre><code># Make sure to set the correct working directory

# Add a new chapter "R for dummies". This will open an R Markdown file.
author_chapter("R_for_dummies")
</code></pre>

<h5>
<a id="user-content-my-chapter-is-finished-how-do-i-upload-it-to-the-datacamp-platform" class="anchor" href="#my-chapter-is-finished-how-do-i-upload-it-to-the-datacamp-platform" aria-hidden="true"><span class="octicon octicon-link"></span></a>My chapter is finished, how do I upload it to the DataCamp platform?</h5>

<p>If you already placed your course on the online DataCamp platform, you can upload a new or existing chapter via the function <code>upload_chapter("chapter_name")</code>. If your course is not yet on the online DataCamp platform, you first need to run the command <code>upload_course()</code>.</p>

<pre><code># Upload the chapter "R for dummies"
upload_chapter("R_for_dummies.Rmd")
</code></pre>

<h5>
<a id="user-content-i-want-to-remove-a-chapter-in-my-course-can-i-do-this" class="anchor" href="#i-want-to-remove-a-chapter-in-my-course-can-i-do-this" aria-hidden="true"><span class="octicon octicon-link"></span></a>I want to remove a chapter in my course. Can I do this?</h5>

<p>Yes you can! To remove chapters out of a course, you need to go to your <code>course.yml</code> file. Here you will see a list of all the chapters included in your course (in this case two), with their chapter ID: </p>

<pre><code># e.g. course.yml
id: 314
title: Hello world
author_field: Your name
description: A first course
chapters:
  - chapter1.Rmd: 753
  - chapter2.Rmd: 760
</code></pre>

<p>To remove a chapter from a course, you just need to remove the corresponding chapter name and ID in the Yaml file. So suppose you want to remove <code>chapter1.Rmd</code> from the course, your Yaml file needs to look like this:</p>

<pre><code># e.g. course.yml
id: 314
title: Hello world
author_field: Your name
description: A first course
chapters:
  - chapter2.Rmd: 760
</code></pre>

<p>To make these changes visible on the DataCamp platform, you need to re-upload the course via <code>upload_course()</code>.</p>

<p><i>Note that removing the first chapter from the Yaml file does not remove the R Markdown file <code>chapter1.Rmd</code> in your course map. So chapter content will not get lost, and you can always re-add the chapter in a later phase.  </i></p>

<h5>
<a id="user-content-how-can-i-add-exercises-to-a-chapter" class="anchor" href="#how-can-i-add-exercises-to-a-chapter" aria-hidden="true"><span class="octicon octicon-link"></span></a>How can I add exercises to a chapter?</h5>

<p>Adding exercises to a chapter is easy. In every R Markdown file of a chapter, the start of a new exercise is indicated by <code>---</code>, followed by the different components of an exercise: <code>## Exercise Title</code>, <code>*** =instructions</code>, <code>*** =hint</code>, <code>*** =pre_exercise_code</code>, <code>*** =solution</code>, <code>*** =sample_code</code>, and <code>*** =sct</code>. </p>

<p>You have an R Markdown with the following exercise (for clarity, only the components <code>## Exercise Title</code> and <code>*** =sct</code> are displayed):</p>

<pre><code>## First Exercise
.
.
.  
*** =sct
</code></pre>

<p>Now you want to add a second exercise. You simply do this by adding the <code>---</code> sign after the <code>*** =sct</code> component of the first exercise:</p>

<pre><code>## First Exercise
.
.
.
.
*** =sct

# AFTER SCT OF EXERCISE ONE, PLACE THE `---` SIGN TO INDICATE START EXERCISE TWO:
---
## Second Exercise
.
.
.
.
*** =sct
</code></pre>

<h5>
<a id="user-content-i-need-to-delete-one-of-the-exercises-in-my-chapter-how-can-i-do-this" class="anchor" href="#i-need-to-delete-one-of-the-exercises-in-my-chapter-how-can-i-do-this" aria-hidden="true"><span class="octicon octicon-link"></span></a>I need to delete one of the exercises in my chapter. How can I do this?</h5>

<p>To delete an exercise from a chapter, you need to delete all the related components of this exercise. The following R Markdown has two exercises and you would like to delete the second one: </p>

<pre><code>## First Exercise
.
.
.
.
*** =sct
---
## Second Exercise
.
.
.
.
*** =sct
</code></pre>

<p>Now just delete all components of the second exercise. Make sure to delete the <code>---</code> sign that indicated the start of the second exercise: </p>

<pre><code>## First Exercise
.
.
.
.
*** =sct
</code></pre>

<h6>
<a id="user-content-if-you-still-did-not-find-the-answer-that-you-were-looking-for-just-send-an-e-mail-to-teachdatacampcom" class="anchor" href="#if-you-still-did-not-find-the-answer-that-you-were-looking-for-just-send-an-e-mail-to-teachdatacampcom" aria-hidden="true"><span class="octicon octicon-link"></span></a>If you still did not find the answer that you were looking for, just send an e-mail to <b><a href="mailto:teach@datacamp.com">teach@datacamp.com</a></b>.</h6>
</article>
  </div>

</div>

<a href="#jump-to-line" rel="facebox[.linejump]" data-hotkey="l" style="display:none">Jump to Line</a>
<div id="jump-to-line" style="display:none">
  <form accept-charset="UTF-8" class="js-jump-to-line-form">
    <input class="linejump-input js-jump-to-line-field" type="text" placeholder="Jump to line&hellip;" autofocus>
    <button type="submit" class="btn">Go</button>
  </form>
</div>

        </div>

      </div><!-- /.repo-container -->
      <div class="modal-backdrop"></div>
    </div><!-- /.container -->
  </div><!-- /.site -->


    </div><!-- /.wrapper -->

      <div class="container">
  <div class="site-footer" role="contentinfo">
    <ul class="site-footer-links right">
        <li><a href="https://status.github.com/" data-ga-click="Footer, go to status, text:status">Status</a></li>
      <li><a href="https://developer.github.com" data-ga-click="Footer, go to api, text:api">API</a></li>
      <li><a href="https://training.github.com" data-ga-click="Footer, go to training, text:training">Training</a></li>
      <li><a href="https://shop.github.com" data-ga-click="Footer, go to shop, text:shop">Shop</a></li>
        <li><a href="https://github.com/blog" data-ga-click="Footer, go to blog, text:blog">Blog</a></li>
        <li><a href="https://github.com/about" data-ga-click="Footer, go to about, text:about">About</a></li>

    </ul>

    <a href="https://github.com" aria-label="Homepage">
      <span class="mega-octicon octicon-mark-github" title="GitHub"></span>
</a>
    <ul class="site-footer-links">
      <li>&copy; 2015 <span title="0.04125s from github-fe122-cp1-prd.iad.github.net">GitHub</span>, Inc.</li>
        <li><a href="https://github.com/site/terms" data-ga-click="Footer, go to terms, text:terms">Terms</a></li>
        <li><a href="https://github.com/site/privacy" data-ga-click="Footer, go to privacy, text:privacy">Privacy</a></li>
        <li><a href="https://github.com/security" data-ga-click="Footer, go to security, text:security">Security</a></li>
        <li><a href="https://github.com/contact" data-ga-click="Footer, go to contact, text:contact">Contact</a></li>
    </ul>
  </div>
</div>


    <div class="fullscreen-overlay js-fullscreen-overlay" id="fullscreen_overlay">
  <div class="fullscreen-container js-suggester-container">
    <div class="textarea-wrap">
      <textarea name="fullscreen-contents" id="fullscreen-contents" class="fullscreen-contents js-fullscreen-contents" placeholder=""></textarea>
      <div class="suggester-container">
        <div class="suggester fullscreen-suggester js-suggester js-navigation-container"></div>
      </div>
    </div>
  </div>
  <div class="fullscreen-sidebar">
    <a href="#" class="exit-fullscreen js-exit-fullscreen tooltipped tooltipped-w" aria-label="Exit Zen Mode">
      <span class="mega-octicon octicon-screen-normal"></span>
    </a>
    <a href="#" class="theme-switcher js-theme-switcher tooltipped tooltipped-w"
      aria-label="Switch themes">
      <span class="octicon octicon-color-mode"></span>
    </a>
  </div>
</div>



    
    

    <div id="ajax-error-message" class="flash flash-error">
      <span class="octicon octicon-alert"></span>
      <a href="#" class="octicon octicon-x flash-close js-ajax-error-dismiss" aria-label="Dismiss error"></a>
      Something went wrong with that request. Please try again.
    </div>


      <script crossorigin="anonymous" src="https://assets-cdn.github.com/assets/frameworks-d22b59d0085e83b7549ba4341ec9e68f80c2f29c8e49213ee182003dc8d568c6.js"></script>
      <script async="async" crossorigin="anonymous" src="https://assets-cdn.github.com/assets/github-0bc0f45c838b5d9d25bc071d2a4b0abe759a093392087dce55cd2caa00ea4f36.js"></script>
      
      

  </body>
</html>

